Measuring changes in Plasmodium falciparum census population size in response to sequential malaria control interventions

Author response:

The following is the authors’ response to the previous reviews

Reviewer #2 (Public review):

In this manuscript, Tiedje and colleagues longitudinally track changes in parasite numbers across four time points as a way of assessing the effect of malaria control interventions in Ghana. Some of the study results have been reported previously, and in this publication, the authors focus on age-stratification of the results. Malaria prevalence was lower in all age groups after IRS. Follow-up with SMC, however, maintained lower parasite prevalence in the targeted age group but not the population as a whole. Additionally, they observe that diversity measures rebound more slowly than prevalence measures. This adds to a growing literature that demonstrates the relevance of asymptomatic reservoirs.

Strengths:

Overall, I found these results clear, convincing, and well-presented. There is growing interest in developing an expanded toolkit for genomic epidemiology in malaria, and detecting changes in transmission intensity is one major application. As the authors summarize, there is no one-size-fits-all approach, and the Bayesian MOIvar estimate developed here has the potential to complement currently used methods, particularly in regions with high diversity/transmission. I find its extension to a calculation of absolute parasite numbers appealing as this could serve as both a conceptually straightforward and biologically meaningful metric.

We thank the reviewer for this positive review of our results and approach.

Weaknesses:

While I understand the conceptual importance of distinguishing among parasite prevalence, mean MOI, and absolute parasite number, I am not fully convinced by this manuscript's implementation of "census population size".

This reviewer remains unconvinced of the use of the term “census population size”. This appears to be due to the dependence of the term on sample size rather than representing a count of a whole population. To give context to our use we are clear in the study presented that the term describes a count of the parasite “strains” in an age-specific sample of a human population in a specified location undergoing malaria interventions.

They have suggested instead using “sample parasite count”. We argue that this definition is too specific and less applicable when we extrapolate the same concept to a different denominator, such as the population in a given area. Importantly, our ecological use of a census allows us to count the appearance of the same strain more than once should this occur in different people.

The authors reference the population genetic literature, but within the context of that field, "census population size" refers to the total population size (which, if not formally counted, can be extrapolated) as opposed to "effective population" size, which accounts for a multitude of demographic factors. There is often interesting biology to be gleaned from the magnitude of difference between N and Ne.

As stated in the introduction we have been explicit in saying that we are not using a population genetic framework. Exploration of N and Ne in population genetics has merit. How this is reconciled when using a “strain” definition and not neutral markers would need to be assessed.

In this manuscript, however, "census population size" is used to describe the number of distinct parasites detected within a sample, not a population. As a result, the counts do not have an immediate population genetic interpretation and cannot be directly compared to Ne. This doesn't negate their usefulness but does complicate the use of a standard population genetic term.

We are clear we are defining a census of parasite strains in an age-specific sample of a population living in two catchment areas of Bongo District. We appreciate the concern of the reviewer and have now further edited the relevant paragraphs in both the Introduction (Lines 75-80) and the Discussion (Lines 501-506) to make very clear the dependence of the reported quantity on sample size, but also its feasible extrapolation consistent with the census of a population.

In contrast, I think that sample parasite count will be most useful in an epidemiological context, where the total number of sampled parasites can be contrasted with other metrics to help us better understand how parasites are divided across hosts, space and time. However, for this use, I find it problematic that the metric does not appear to correct for variations in participant number. For instance, in this study, participant numbers especially varied across time for 1-5 year-olds (N=356, 216, 405, and 354 in 2012, 2014, 2015, and 2017 respectively).

The reviewer has made an important point that for the purpose of comparisons across the four surveys or study time points (i.e., 2012, 2014, 2015, and 2017), we should "normalize" the number of individuals considered for the calculation of the "census population size". Given that this quantity is a sum of the estimated MOIvar,, we need to have constant numbers for its values to be compared across the surveys, within age group and the whole population. This is needed not only to get around the issue of the drop in 1-5 year olds surveyed in 2014 but to also stabilize the total number of individuals for the whole sample and for specific age groups. One way to do this is to use the smaller sample size for each age group across time, and to use that value to resample repeatedly for that number of individuals for surveys where we have a larger sample size. This has now been updated included in the manuscript as described in the Materials and Methods (Lines 329-341) and in the Results (Lines 415-430; see updated Figure 4 and Table supplement 7). This correction produces very similar results to those we had presented before (see updated Figure 4 and Table supplement 7).

As stated in our previous response we have used participant number in an interrupted time series where the population was sampled by age to look at age-specific effects of sequential interventions IRS and SMC. As shown in Table supplement 1 of the 16 age-specific samples of the total population, we have sampled very similar proportions of the population by age group across the four surveys. The only exception was the 1-5 year-old age group during the survey in 2014. We are happy to provide additional details to further clarify the lower number (or percentage) of 1-5 year olds (based on the total number of participants per survey) in 2014 (~12%; N = 216) compared to the other surveys conducted 2012, 2015, and 2017 (~18-20%; N = 356, 405, and 354, respectively). Please see Table supplement 1 for the total number of participants surveyed in each of the four surveys (i.e., 2012, 2014, 2015, and 2017).

This sample size variability is accounted for with other metrics like mean MOI.

We agree that mean MOI by age presents a way forward with variable samples to scale up. Please see updated Figure supplement 8.

In sum, while the manuscript opens up an interesting discussion, I'm left with an incomplete understanding of the robustness and interpretability of the new proposed metric.”

We thank you for your opinion. We have further edited the manuscript to make clear our choice of the term and the issue of sample size. We believe the proposed terminology is meaningful as explained above.

Reviewer #3 (Public review):

Summary

The manuscript coins a term "the census population size" which they define from the diversity of malaria parasites observed in the human community. They use it to explore changes in parasite diversity in more than 2000 people in Ghana following different control interventions.

Strengths:

This is a good demonstration of how genetic information can be used to augment routinely recorded epidemiological and entomological data to understand the dynamics of malaria and how it is controlled. The genetic information does add to our understanding, though by how much is currently unclear (in this setting it says the same thing as age stratified parasite prevalence), and its relevance moving forward will depend on the practicalities and cost of the data collection and analysis. Nevertheless, this is a great dataset with good analysis and a good attempt to understand more about what is going on in the parasite population.

Thank you to the reviewer for their supportive assessment of our research.

Weaknesses

None

Reviewer #3 (Recommendations for the authors):

New figure supplement 8 - x-axis says percentage but goes between 0-1, so is a proportion

We thank the reviewer for bringing this to our attention. We have amended the x-axis labels accordingly for Figure supplement 8.

Changes from the earlier version to the revised version

Abstract

Abstract

Here we introduce a new endpoint “census population size” to evaluate the epidemiology and control of Plasmodium falciparum infections, where the parasite, rather than the infected human host, is the unit of measurement. To calculate census population size, we rely on a definition of parasite variation known as multiplicity of infection (MOIvar), based on the hyper-diversity of the var multigene family. We present a Bayesian approach to estimate MOIvar from sequencing and counting the number of unique DBLα tags (or DBLα types) of var genes, and derive from it census population size by summation of MOIvar in the human population. We track changes in this parasite population size and structure through sequential malaria interventions by indoor residual spraying (IRS) and seasonal malaria chemoprevention (SMC) from 2012 to 2017 in an area of high-seasonal malaria transmission in northern Ghana. Following IRS, which reduced transmission intensity by > 90% and decreased parasite prevalence by ∼40-50%, significant reductions in var diversity, MOIvar, and population size were observed in ∼2,000 humans across all ages. These changes, consistent with the loss of diverse parasite genomes, were short lived and 32-months after IRS was discontinued and SMC was introduced, var diversity and population size rebounded in all age groups except for the younger children (1-5 years) targeted by SMC. Despite major perturbations from IRS and SMC interventions, the parasite population remained very large and retained the var population genetic characteristics of a high-transmission system (high var diversity; low var repertoire similarity) demonstrating the resilience of P. falciparum to short-term interventions in high-burden countries of sub-Saharan Africa.

Introduction

Malaria in high-transmission endemic areas of sub-Saharan Africa (SSA) is characterised by vast diversity of the Plasmodium falciparum parasites from the perspective of antigenic variation (Chen et al., 2011; Day et al., 2017; Otto et al., 2019; Ruybal-Pesántez et al., 2022, 2017). As with other hosts of hyper-variable pathogens (Futse et al., 2008), children experiencing clinical episodes of malaria, eventually become immune to disease but not to infection. This results in a large reservoir of chronic asymptomatic infections, in hosts of all ages, sustaining transmission to mosquitos. Given the goal of malaria eradication by 2050, it is therefore of interest to examine how the parasite population changes following perturbation by major intervention efforts, both in terms of its size and underlying population genetics.

Earlier version: So, what do we mean by the parasite population size in the case of P. falciparum and how do we measure it? Parasite prevalence, detected by microscopy or more sensitive molecular diagnostics (e.g., PCR), describes the proportion of infected human hosts. Studies of P. falciparum genetic diversity have shown that the majority of people in high-transmission endemic areas harbour diverse multiclonal infections measured as the complexity or multiplicity of infection (MOI) (e.g., (Anderson et al., 2000; Paul et al., 1995; Smith et al., 1999; Sumner et al., 2021)) with complex population dynamics (Bruce et al., 2000; Farnert et al., 1997). These genetic data indicate much larger parasite population sizes than observed by prevalence of infection alone. Thus, from an ecological perspective, we can consider a human host as a patch carrying a number of “antigenically-distinct infections” of P. falciparum. The sum of these antigenically-distinct infections over all hosts provides us with a census of the parasite count of relevance to monitoring and evaluating malaria interventions. We refer to this census population size hereafter simply as population size but make clear that this measure is distinct from effective population size (Ne) as measured by neutral variation.

Revised version: So, what do we mean by the parasite population size in the case of P. falciparum and how do we measure it? Parasite prevalence, detected by microscopy or more sensitive molecular diagnostics (e.g., PCR), describes the proportion of infected human hosts. Studies of P. falciparum genetic diversity have shown that the majority of people in high-transmission endemic areas harbour diverse multiclonal infections measured as the complexity or multiplicity of infection (MOI) (e.g., (Anderson et al., 2000; Paul et al., 1995; Smith et al., 1999; Sumner et al., 2021)) with complex population dynamics (Bruce et al., 2000; Farnert et al., 1997). These genetic data indicate much larger parasite population sizes than observed by prevalence of infection alone. Thus, from an ecological perspective, we can consider a human host as a patch carrying a number of “antigenically-distinct infections” of P. falciparum. The sum of these antigenically-distinct infections over all sampled hosts provides us with a census of the parasite count of relevance to monitoring and evaluating malaria interventions. We refer to this census population size hereafter simply as population size but make clear that this measure is distinct from effective population size (Ne) as measured by neutral variation. This count can be scaled from the host sample to the larger denominator of a host population in the area of interest.

Diversity of P. falciparum single copy surface antigen genes such as circumsporozoite protein (csp), merozoite surface protein 1 (msp1) or 2 (msp2), and apical membrane antigen 1 (ama1) have each been widely used to measure MOI (e.g., (Falk et al., 2006; Lerch et al., 2017; Nelson et al., 2019)). They have become part of newer genetic panels (e.g., Paragon v1 (Tessema et al., 2020) and AMPLseq v1 (LaVerriere et al., 2022)) specifically for MOI determination. Typically, MOI is reported as the maximum number of alleles or single locus haplotypes present at the most diverse of these antigen-encoding loci rather than the number of unique multilocus haplotypes of these genes combined, as it is challenging to accurately reconstruct or phase these haplotypes in hosts with an MOI > 3 (Lerch et al., 2019). Each of these genes is under balancing selection with a few geographically-common haplotypes and many very rare haplotypes in moderate- to high-transmission settings (Markwalter et al., 2022; Sumner et al., 2021). Where there is a high probability of co-occurrence of two or more common single locus haplotypes in a host, genotyping each of these single copy antigen genes alone will underestimate MOI. Single nucleotide polymorphism (SNP) panels have been used to define the presence of multiclonal infections with limited reliability to estimate MOI for highly complex infections, typical in high transmission, even with the use of computational methods (Labbé et al., 2023).

As an alternative to genotyping single copy antigen genes and biallelic SNP panels to estimate MOI, we have proposed the use of a fingerprinting methodology known as varcoding to genotype the hyper-diverse var multigene family (∼50-60 var genes per haploid genome). This method employs a ∼450bp region of a var gene, known as a DBLα tag encoding the immunogenic Duffy-binding-like alpha (DBLα) domain of PfEMP1 (Plasmodium falciparum erythrocyte membrane protein 1), the major surface antigen of the blood stages (Zhang and Deitsch, 2022). Bioinformatic analyses of a large database of exon 1 sequences of var genes showed a predominantly 1-to-1 DBLα-var relationship, such that each DBLα tag typically represents a unique var gene, especially in high transmission (Tan et al., 2023). The extensive diversity of DBLα tags together with the very low percentage of var genes shared between parasites (Chen et al., 2011; Day et al., 2017; Ruybal-Pesántez et al., 2022, 2017) facilitate measuring MOI by amplifying, pooling, sequencing, and counting the number of unique DBLα tags (or DBLα types) in a host (Ruybal-Pesántez et al., 2022; Tiedje et al., 2022). From a single PCR with degenerate primers and amplicon sequencing, the method specifically counts the most diverse DBLα types, designated non-upsA, per infection to arrive at a metric we call MOIvar. It is not based on assigning haplotypes but exploits the fact that var repertoires are non-overlapping, especially in high transmission. Instead, it assumes a set number of non-upsA types per genome based on repeated sampling of 3D7 control isolates accounting for PCR sampling errors to calculate MOIvar (Ghansah et al., 2023; Ruybal-Pesántez et al., 2022; Tiedje et al., 2022). Consequently, rather than looking at the diversity of a single copy antigen-encoding gene like csp, msp2, or ama1 to calculate MOI, by varcoding we are looking at sets of up to 45 non-upsA DBLα types per genome. Prior work has shown that varcoding is more sensitive to measure MOI in high transmission where there is an extremely high prevalence of multiclonal infections that cannot be accurately phased with either biallelic SNP panels (Ghansah et al., 2023; Labbé et al., 2023; Tessema et al., 2020; Watson et al., 2021) or combinations of single copy antigen genes (Sumner et al., 2021).

Here we report an investigation of changes in parasite census population size and structure through two sequential malaria control interventions between 2012 and 2017 in Bongo District located in the Upper East Region of Ghana, one of the 12 highest burden countries in Africa (World Health Organization, 2022). We present a novel Bayesian modification to the published varcoding approach (Ghansah et al., 2023; Ruybal-Pesántez et al., 2022; Tiedje et al., 2022) that takes into account under-sampling of non-upsA DBLα types in an isolate to estimate MOIvar (Ruybal-Pesántez et al., 2022; Tiedje et al., 2022) and therefore population size. We document P. falciparum prevalence as well as var diversity and population structure from baseline in 2012 through a major perturbation by a short-term indoor residual spraying (IRS) campaign managed under operational conditions, which reduced transmission intensity by > 90% as measured by the entomological inoculation rate (EIR) and decreased parasite prevalence by ∼40-50% (Tiedje et al., 2022). Next, we followed what happened to parasite population size more than two years after the IRS intervention was discontinued and seasonal malaria chemoprevention (SMC) was introduced for children between the ages of 3-59 months (i.e., < 5 years) (Wagman et al., 2018). Detectable changes in parasite population size were seen as a consequence of the IRS intervention but this quantity rapidly rebounded 32-months after the intervention ceased. Overall, throughout the IRS, SMC, and subsequent rebound, the parasite population in humans remained large in size and retained the var population genetic characteristics of high transmission (i.e., high var diversity, low var repertoire overlap), demonstrating the overall resilience of the species to survive significant short-term perturbations.

Materials and methods

Human subject ethical approval

Earlier version: The study was reviewed/approved by the ethics committees at the Navrongo Health Research Centre (Ghana), Noguchi Memorial Institute for Medical Research (Ghana), The University of Melbourne (Australia), and The University of Chicago (United States). Details on the study area, study population, inclusion/exclusion criteria, and data collection procedures have been previously described (Tiedje et al., 2022, 2017).

Revised version: The study was reviewed/approved by the ethics committees at the Navrongo Health Research Centre (Ghana), Noguchi Memorial Institute for Medical Research (Ghana), The University of Melbourne (Australia), The University of Chicago (United States), and New York University (United States). Details on the study area, study population, inclusion/exclusion criteria, and data collection procedures have been previously described (Tiedje et al., 2022, 2017).

Study design and sample collection

Earlier version: Using an interrupted time-series study design, four age-stratified cross-sectional surveys of ∼2,000 participants per survey were undertaken to investigate the impacts of IRS and SMC in combination with long-lasting insecticidal nets (LLINs) impregnated with pyrethroids under operational conditions on the asymptomatic P. falciparum reservoir from two proximal catchment areas (i.e., Vea/Gowrie and Soe, with a sampling area of ∼60 km2) in Bongo District, Ghana (Tiedje et al., 2022, 2017) (Table supplement 1). Bongo District, located in the Upper East Region, is categorized as high-seasonal malaria transmission based on the World Health Organization’s (WHO) “A Framework for Malaria Elimination” (WHO/GMP, 2017) where P. falciparum prevalence was ≥ 35% at baseline in 2012 (Tiedje et al., 2022, 2017). These ∼2,000 participants of all ages (1-97 years) represent ∼15% of the total population that reside in these two catchment areas in Bongo (Tiedje et al., 2017). The four cross-sectional surveys were completed at the end of the wet season (i.e., high-transmission season) and the study can be separated into four distinct study time points: (1) October 2012 (Survey 1) prior to the IRS and SMC (i.e., baseline), (2) October 2014 (Survey 2) two months after the second round of IRS (Actellic 50EC), (3) October 2015 (Survey 3) seven months after the third round of IRS using a long-acting non-pyrethroid insecticide (Actellic 300CS) (Gogue et al., 2020; US President’s Malaria Initiative Africa IRS (AIRS) Project, 2016), and finally (4) October 2017 (Survey 4) 32-months after the discontinuation of IRS, but during the deployment of SMC to all children 3-59 months (i.e., < 5 years) (Figure 1). LLINs (i.e., PermaNet 2.0, Olyset, or DawaPlus 2.0) were mass distributed in Bongo District by the National Malaria Elimination Programme (NMEP)/Ghana Health Service (GHS) between 2010-2012 and again in 2016 following the discontinuation of IRS (Gogue et al., 2020; Tiedje et al., 2022; US Agency for International Development (USAID) Global Health Supply Chain Program, 2020). In addition, to maintain high coverage of LLINs between these campaigns continuous distribution was undertaken using routine services (i.e., antenatal clinics, school distributions, immunization visits, etc.) (Tiedje et al., 2022). Over the course of this study self-reported LLIN usage the previous night remained high across all age groups, from 89.1% in 2012 (pre-IRS), 83.5% in 2014 (during IRS), 90.6% in 2015 (post-IRS), to 96.8% in 2017 (SMC). Details on the study area, study population, and data collection procedures have been previously described (Tiedje et al., 2022, 2017).

Revised version: Using an interrupted time-series study design, four age-stratified cross-sectional surveys of ∼2,000 participants per survey were undertaken to investigate the impacts of IRS and SMC in combination with long-lasting insecticidal nets (LLINs) impregnated with pyrethroids under operational conditions on the asymptomatic P. falciparum reservoir from two proximal catchment areas (i.e., Vea/Gowrie and Soe, with a sampling area of ∼60 km2) in Bongo District, Ghana (hereinafter referred to collectively as “Bongo”; (Tiedje et al., 2022, 2017)) (Table supplement 1). Bongo District, located in the Upper East Region, is categorized as high-seasonal malaria transmission based on the World Health Organization’s (WHO) “A Framework for Malaria Elimination” (WHO/GMP, 2017) where P. falciparum prevalence was ≥ 35% at baseline in 2012 (Tiedje et al., 2022, 2017). These ∼2,000 participants of all ages (1-97 years) represent ∼15% of the total population that reside in these two catchment areas in Bongo (Tiedje et al., 2017). The four cross-sectional surveys were completed at the end of the wet season (i.e., high-transmission season) and the study can be separated into four distinct study time points: (1) October 2012 (Survey 1) prior to the IRS and SMC (i.e., baseline), (2) October 2014 (Survey 2) two months after the second round of IRS (Actellic 50EC), (3) October 2015 (Survey 3) seven months after the third round of IRS using a long-acting non-pyrethroid insecticide (Actellic 300CS) (Gogue et al., 2020; US President’s Malaria Initiative Africa IRS (AIRS) Project, 2016), and finally (4) October 2017 (Survey 4) 32-months after the discontinuation of IRS, but during the deployment of SMC to all children 3-59 months (i.e., < 5 years) (Figure 1). LLINs (i.e., PermaNet 2.0, Olyset, or DawaPlus 2.0) were mass distributed in Bongo District by the National Malaria Elimination Programme (NMEP)/Ghana Health Service (GHS) between 2010-2012 and again in 2016 following the discontinuation of IRS (Gogue et al., 2020; Tiedje et al., 2022; US Agency for International Development (USAID) Global Health Supply Chain Program, 2020). In addition, to maintain high coverage of LLINs between these campaigns continuous distribution was undertaken using routine services (i.e., antenatal clinics, school distributions, immunization visits, etc.) (Tiedje et al., 2022). Over the course of this study self-reported LLIN usage the previous night remained high across all age groups, from 89.1% in 2012 (pre-IRS), 83.5% in 2014 (during IRS), 90.6% in 2015 (post-IRS), to 96.8% in 2017 (SMC). Details on the study area, study population, and data collection procedures have been previously described (Tiedje et al., 2022, 2017).

Study design and changes in the prevalence of microscopic P. falciparum infection following the IRS and SMC interventions in Bongo, Ghana.

(A) Four age-stratified cross-sectional surveys of ∼2,000 participants per survey were conducted in Bongo, Ghana at the end of the wet seasons in October 2012 (Survey 1, baseline pre-IRS, red), October 2014 (Survey 2, during IRS, orange), October 2015 (Survey 3, post-IRS, green), and October 2017 (Survey 4, SMC, purple) (Table supplement 1) (see Materials and Methods). The three rounds of IRS (grey areas) were implemented between 2013 and 2015 (Tiedje et al., 2022). SMC was distributed to all children < 5 years of age during the wet seasons in 2016 (two rounds between August – September 2016) and 2017 (four rounds between September – December 2017) (Gogue et al., 2020). Both IRS and SMC were implemented against a background of widespread LLIN usage (Tiedje et al., 2022). Prevalence of microscopic P. falciparum infections (%) in the (B) study population and (C) for all age groups (years) in each survey. Error bars represent the upper and lower limits of the 95% confidence interval (CI) calculated using the Wald Interval.

Details of the IRS and SMC interventions

Earlier version: Starting in 2013, the AngloGold Ashanti Malaria Control Programme (AGAMal) in a public-private partnership with the Global Fund, scaled up IRS across all of the Upper East Region of northern Ghana (Gogue et al., 2020). As part of this initiative, three rounds of IRS with organophosphate formulations (i.e., non-pyrethroid) were rolled out prior to the start of the wet season between 2013 and 2015 (Figure 1A) in Bongo District (Tiedje et al., 2022). Based on AGAMal’s operational reports, IRS coverage in Bongo District was 91.8% in Round 1, 95.6% in Round 2, and finally 96.6% in Round 3 (AGAMal, personal communication). To monitor the impact of the IRS on the local vector population, monthly entomological collections were undertaken between February 2013 and September 2015 (Tiedje et al., 2022). Using these surveys, we observed that the monthly entomological inoculation rate (EIR) (infective bites/person/month (ib/p/m)), a measure of local transmission intensity, declined by > 90% at the peak of the wet season between August 2013 (pre-IRS) (EIR = 5.3 ib/p/m) and August 2015 (post-IRS) (EIR = 0.4 ib/p/m) (Tiedje et al., 2022). Following the IRS, SMC was rolled out in the Upper East Region by the NMEP)/(GHS) starting in 2016 (Gogue et al., 2020) (Figure 1A). SMC is the intermittent administration of full treatment courses of an antimalarial to children between the ages of 3-59 months (i.e., < 5 years) (WHO, 2012). Like other countries, the SMC drug combination of choice in Ghana is amodiaquine plus sulfadoxine-pyrimethamine, which is administered at monthly intervals during the high-transmission season (i.e., wet season) (WHO, 2012). The goal of this age-targeted intervention is to both clear current malaria infections and prevent malarial illness by maintaining a therapeutic concentration of an antimalarial in blood over the period of greatest risk (i.e., high-transmission season). Reported SMC coverage in Bongo District was 92.6% in 2016 (two rounds between August – September 2016) and 94.6% in 2017 (four rounds between September – December 2017) (Gogue et al., 2020).

Revised version: Starting in 2013, the AngloGold Ashanti Malaria Control Programme (AGAMal) in a public-private partnership with the Global Fund, scaled up IRS across all of the Upper East Region of northern Ghana (Gogue et al., 2020). As part of this initiative, three rounds of IRS with organophosphate formulations (i.e., non-pyrethroid) were rolled out prior to the start of the wet season between 2013 and 2015 (Figure 1A) in Bongo District (Tiedje et al., 2022). Based on AGAMal’s operational reports, IRS coverage in Bongo District was 91.8% in Round 1, 95.6% in Round 2, and finally 96.6% in Round 3 (AGAMal, personal communication). To monitor the impact of the IRS on the local vector population, monthly entomological collections were undertaken between February 2013 and September 2015 (Tiedje et al., 2022). Using these surveys, we observed that the monthly entomological inoculation rate (EIR) (infective bites/person/month (ib/p/m)), a measure of local transmission intensity, declined by > 90% at the peak of the wet season between August 2013 (pre-IRS) (EIR = 5.3 ib/p/m) and August 2015 (post-IRS) (EIR = 0.4 ib/p/m) (Tiedje et al., 2022). Following the IRS, SMC was rolled out in the Upper East Region by the NMEP/GHS starting in 2016 (Gogue et al., 2020) (Figure 1A). SMC is the intermittent administration of a curative dose of an antimalarial to children between the ages of 3-59 months (i.e., < 5 years) (WHO, 2012). Like other countries, the SMC drug combination of choice in Ghana is sulfadoxine-pyrimethamine plus amodiaquine, which is administered at monthly intervals (i.e., ∼28-30 days apart) during the high-transmission season (i.e., wet season) (WHO, 2012). The goal of this age-targeted intervention is to both clear current infections and prevent malarial illness by maintaining a therapeutic concentration of an antimalarial in the blood over the period of greatest risk (i.e., high-transmission season). Reported SMC coverage in Bongo District was 92.6% in 2016 (two cycles between August – September 2016) and 94.6% in 2017 (four cycles between July – October 2017) (Gogue et al., 2020) (NMEP/GHS, personal communication).

Var genotyping and sequence analysis

Earlier version: Genomic DNA was extracted from the dried blood spots for all participants with a confirmed microscopic asymptomatic P. falciparum infection (i.e., isolate) (2,572 isolates, Table supplement 2) using the QIAamp DNA mini kit (QIAGEN, USA) as previously described (Tiedje et al., 2017). For var genotyping or varcoding, the sequence region within the var genes encoding the DBLα domains of PfEMP1 were amplified using a single-step PCR, pooled, and sequenced on an Illumina platform using the MiSeq 2×300bp paired-end protocol (Supplemental Methods, Figure supplement 1). The raw sequence data was then processed using our previously published customized bioinformatic pipelines (Supplemental Methods, Figure supplement 2). For additional information on the use of these bioinformatic pipelines a detailed tutorial is available (https://github.com/UniMelb-Day-Lab/tutorialDBLalpha).

Revised version: Genomic DNA was extracted from the dried blood spots for all participants with a confirmed microscopic asymptomatic P. falciparum infection (i.e., isolate) (2,572 isolates, Table supplement 2) using the QIAamp DNA mini kit (QIAGEN, USA) as previously described (Tiedje et al., 2017). For var genotyping or varcoding, the sequence region within the var genes encoding the DBLα domains of PfEMP1 were amplified using a single-step PCR, pooled, and sequenced on an Illumina platform using the MiSeq Reagent Kit v3 (600 cycle; 2×300bp paired-end) (Supplemental Methods, Figure supplement 1). The raw sequence data was then processed using our previously published customized bioinformatic pipelines (Supplemental Methods, Figure supplement 2). For additional information on the use of these bioinformatic pipelines a detailed tutorial is available (https://github.com/UniMelb-Day-Lab/tutorialDBLalpha).

DBLα sequencing data was obtained from 2,397 P. falciparum isolates (93.2%) (Table supplement 2). This genotyping success was acceptable given that we were working with low-density asymptomatic infections (Table supplement 1). Using a cut-off of ≥ 20 DBLα types, DBLα sequencing data was obtained from 2,099 P. falciparum isolates (81.6%) with a total of 289,049 DBLα sequences and 53,238 unique DBLα types being identified in the study population (Table supplement 2). The median P. falciparum density was ∼4 times higher for isolates with ≥ 20 DBLα types compared to those that gave no or < 20 DBLα types (520 [IQR: 200-1880] parasites/μL vs. 120 [IQR: 40-200] parasites/μL, respectively).

DBLα type diversity

We monitored the impacts of the sequential interventions (i.e., IRS and SMC) on diversity by measuring changes in the population genetics of DBLα types at the population level (i.e., P. falciparum reservoir). Diversity was monitored using two measures, DBLα type richness and DBLα type frequency. Richness was defined as the number of unique DBLα types observed (i.e., DBLα type pool size) in each survey or study time point (i.e., 2012, 2014, 2015, and 2017). DBLα type richness, however, does not provide any information about the relative frequencies of the different DBLα types in the population, as they are all weighted equally whether they are observed once or more frequently (e.g., observed in > 20 isolates per survey). To further examine the impacts of the interventions on DBLα type diversity we also assessed the frequency of each unique DBLα type in 2012, 2014, 2015, and 2017. Here we defined DBLα type frequency as the number of times (i.e., number of isolates) a DBLα type was observed in each survey. Both upsA and non-upsA DBLα type diversities were measured due to their different biological features, chromosomal positions (i.e., subtelomeric regions vs. internal or central regions), as well as population genetics (Falk et al., 2009; Gardner et al., 2002; Jensen et al., 2004; Kaestli et al., 2006; Kalmbach et al., 2010; Kraemer et al., 2007; Kraemer and Smith, 2006; Kyriacou et al., 2006; Lavstsen et al., 2003; Normark et al., 2007; Rottmann et al., 2006; Warimwe et al., 2012, 2009; Zhang and Deitsch, 2022). The proportions of upsA and non-upsA var genes in a repertoire or single genome have been defined as ∼15-20% and ∼80-85%, respectively, based on whole genome sequencing (Rask et al., 2010). The upsA and non-upsA DBLα type proportions were partitioned as expected in our analyses, with the median proportions at the repertoire level being comparable in 2012 (19% upsA and 81% non-upsA), 2014 (22% upsA and 78% non-upsA), 2015 (21% upsA and 79% non-upsA), and 2017 (20% upsA and 80% non-upsA).

Repertoire similarity as defined by pairwise type sharing

To estimate genetic similarity between the DBLα repertoires (i.e., unique DBLα types identified in each isolate) identified from two isolates, pairwise type sharing (PTS) was calculated between all pairs of isolates in each survey as previously described (Barry et al., 2007). PTS, analogous to the Sørensen Index, is a similarity statistic to evaluate the proportion of DBLα types shared between two isolate repertoires (i.e., DBLα repertoire overlap) and ranges from 0 (i.e., no DBLα repertoire overlap) to 1 (i.e., identical DBLα isolate repertoires), where < 0.50 = unrelated, 0.5 = recent recombinants/siblings, > 0.5 = related, and 1 = clones. PTS is a measure of identity-by-state (IBS) used to assess repertoire similarity between isolates and is not used to infer inheritance from a recent common ancestor (i.e., identity-by-decent (IBD)) (Speed and Balding, 2015).

DBLα isolate repertoire size

For this study we have exploited the unique population structure of non-overlapping DBLα isolate repertoires to estimate isolate MOIvar. To calculate MOIvar, the non-upsA DBLα types were chosen since not only are they more diverse and less conserved between isolate repertoires (i.e., low median PTSnon-upsA ≤ 0.020) compared to the upsA DBLα types, but they have also been shown to have a more specific 1-to-1 relationship with a single var gene than upsA (Tan et al., 2023). The low to non-existent overlap of repertoires enables an estimation of MOI that relies on the number of non-upsA DBLα types sequenced from an individual’s isolate (Ruybal-Pesántez et al., 2022; Tiedje et al., 2022). A constant repertoire size or number of DBLα types in a parasite genome can be used to convert the number of types sequenced in an isolate to estimate MOI (Ruybal-Pesántez et al., 2022; Tiedje et al., 2022). This approach, however, neglects the measurement error in this size introduced by targeted PCR and amplicon sequencing of var genes in an isolate.

Added in the revised version: Bayesian estimation of MOIvar and associated census population size

Earlier version: Bayesian estimation of MOIvar and associated census population size Here, we extend the method to a Bayesian formulation and estimate the posterior distribution for each sampled individual for the probability of different MOI values. From individual posterior distributions, we can then obtain the estimated MOI frequency distribution for the population as a whole. The two pieces of information required for our approach are the measurement error and the prior distribution of MOI. The measurement error is simply the repertoire size distribution, that is, the distribution of the number of non-upsA DBLα types sequenced given MOI = 1, which is empirically available (Figure supplement 3 (Labbé et al., 2023)). We refer to it as P(s | MOI = 1) where s here denotes repertoire size. More generally when MOI ≥ 1, s denotes the number of non-upsA DBLα types sequenced or isolate size. The prior distribution of MOI refers to the belief we have for what the actual MOI distribution might look like at the population level before empirical evidence is taken into consideration. For example, the prior distribution of MOI is likely to centre around a higher value in high-transmission endemic areas than in low-transmission ones.

Revised version: Here, we extend the method to a Bayesian formulation and estimate the posterior distribution for each sampled individual for the probability of different MOI values. From individual posterior distributions, we can then obtain the estimated MOI frequency distribution for the population as a whole. The two pieces of information required for our approach are the measurement error and the prior distribution of MOI. The measurement error is simply the repertoire size distribution, that is, the distribution of the number of non-upsA DBLα types sequenced given MOI = 1, which is empirically available (Figure supplement 3 (Labbé et al., 2023)). We refer to it as P(s | MOI = 1) where s here denotes repertoire size. More generally when MOI ≥ 1, s denotes the number of non-upsA DBLα types sequenced or isolate size. The prior distribution of MOI refers to the belief we have for what the actual MOI distribution might look like at the population level before empirical evidence is taken into consideration. For example, the prior distribution of MOI is likely to centre around a higher value in high-transmission endemic areas than in low-transmission ones.

We can obtain P(s | MOI = m) from the serial convolution of the repertoire size distribution P(s | MOI = 1) and P(s | MOI = m - 1). Starting with the repertoire size distribution given a single infection, we can derive P(s | MOI = m) for m equal to 2,3,…, up until a maximum value of 20 (empirically determined), as follows: where L and U are the lower and upper limit for the repertoire size, 10 and 45 respectively from the empirical repertoire size distribution (Figure supplement 3 (Labbé et al., 2023)).

For simplicity, we begin with a uniform prior. We use Bayes’ rule to derive a posterior distribution of MOI given a certain number of non-upsA DBLα types sequenced from an individual: where k is the maximum value of MOI, here 20, as empirically determined.

To obtain the MOI distribution at the population level, we could either simply pool the maximum a posteriori MOI estimate for each sampled individual, or use a technique called mixture distribution. For the latter, we weighed each posterior MOI distribution for each sampled individual equally and sum over all posterior distributions at the individual level to derive the MOI distribution at the population level: where n is the number of sampled individuals. These two approaches gave similar results for our empirical survey data as determined by the Kolmogorov-Smirnov Test. The obtained distance statistic is close to 0 and the corresponding p-value is non-significant across all surveys, indicating that the two estimates were drawn from the same distribution (Table supplements 3 and 4). Additionally, the difference between the mean MOI values at the population level obtained from the two approaches is small (Table supplements 3 and 4). Given this similarity, we present the results based on pooling the maximum a posteriori MOI estimates for each sampled individual in the main text and include the results based on mixture distribution in the Supplemental Information. Note that we focused on individuals who had confirmed microscopic asymptomatic P. falciparum infections for our MOI estimation.

To examine alternative priors, we considered empirical MOI distributions described in the literature including the Poisson, hyper-Poisson, and negative binomial distributions (Dietz, 1988; Henry, 2020). The hyper-Poisson and negative binomial distributions can capture the overdispersion seen in the empirical distribution of MOI for certain areas and caused by factors such as heterogeneous biting. We therefore focused on a negative binomial distribution and investigated changing its parameters to generate priors with different means spanning a wide range of MOI values (mean MOI within [∼1.5, ∼6.7]), including those seen in high-transmission endemic areas. A uniform prior and a zero-truncated negative binomial distribution with parameters within the range typical of high-transmission endemic regions (higher mean MOI, for example, ∼4.3 versus ∼6.7, with tails for higher MOI values in the range of 10-20) produce similar MOI estimates at the population level (Table supplements 5 and 6). However, when setting the parameter range of the zero-truncated negative binomial to be representative of those in low-transmission endemic regions where the empirical MOI distribution centres around monoclonal infections with the majority of MOIs = 1 or 2 (mean MOI ≈ 1.5, no tail at higher MOI values), the final population-level MOI distribution does deviate more from that based on the aforementioned prior and parameter choices (Table supplement 6). The final individual- and population-level MOI estimates are not sensitive to the specifics of the prior MOI distribution as long as this distribution captures the tail for higher MOI values with above-zero probability. The obtained Kolmogorov-Smirnov Test distance statistics and their corresponding p-values, the Pearson correlation tests and their corresponding p-values, as well as the difference in mean MOI values, for the comparison of the MOI estimates obtained with the different priors are included in Table supplements 5 and 6. Given these comparisons, we provide in our analyses the estimated population MOI distribution using a uniform prior.

Added in the revised version: Adjusting for differences in sampling depth across age groups and surveys in the census population size calculation

Added in the revised version: In our interrupted time-series study design, we sampled the Bongo population by age group to assess age-specific effects of the sequential interventions (i.e., IRS and SMC). Sampling depth by age group was consistent across the four cross-sectional surveys (Table supplement 1), except for 1-5-years-olds in 2014. To account for variation in sampling depth across age groups and surveys, we performed subsampling at the level of the minimum number of individuals in each age group across all surveys. We calculated the total number of var repertoires from a subsample of this number of individuals within each age group for each survey. This approach ensures consistent sample sizes within each age group across all surveys. We summed the var repertoires across age groups to obtain the total var repertoire count for each survey. To calculate the mean and 95% CIs for the number of var repertoires (i.e., census population size), we repeated the subsampling procedure 10,000 times and derived these quantities from the distribution of these repeated subsampling replicates.

Statistical analysis

We used the R v4.3.1 for all data analyses with the collection of R packages in tidyverse being used for data curation along with base, stats, gtsummary, and epiR for the statistical analyses (R Core Team, 2018; Sjoberg, D et al., 2021; Stevenson, 2020; Wickham et al., 2019). Continuous variables are presented as medians with interquartile ranges (IQRs) and discrete variables are presented using the observed or calculated values with the 95% confidence interval (95% CIs) or ± 2 standard deviations (± 2SD). Kaplan-Meier survival curves were generated for the time (i.e., number of surveys) to first event (i.e., when the DBLα type was no longer observed/detected) comparing the upsA and non-upsA DBLα types; p-values were determined using the log-rank test using the R packages survival and survminer (Kassambara et al., 2021; Therneau, 2023). The time interval to first event considered for all survival curves was the number of surveys or year (i.e., 2012, 2014, 2015, and 2017) that each DBLα type was observed and only includes those upsA (N = 2,218) and non-upsA (N = 33,159) DBLα types that were seen at baseline in 2012 (i.e., those DBLα types observed prior to the IRS intervention) (Table supplement 2).

Results

Between 2013 and 2015, three-rounds of IRS with non-pyrethroid insecticides were implemented across all of Bongo District. Coincident with the > 90% decrease in transmission following IRS (Tiedje et al., 2022), the prevalence of microscopic P. falciparum infections compared to the 2012 baseline survey (pre-IRS) declined by 45.2% and 35.7% following the second (2012 to 2014) and the third (2012 to 2015) round of IRS, respectively (Figure 1B, Table supplement 1). These declines in parasite prevalence were observed across all ages, with the greatest impacts being observed on the younger children (1-5 years) who were ∼3 times less likely to have an infection in 2015 (post-IRS) compared to 2012 (pre-IRS) (Figure 1C, Table supplement 1). These reductions were however short-lived and in 2017, 32-months after the discontinuation of IRS, but during SMC, overall P. falciparum prevalence rebounded to 41.2%, or near pre-IRS levels (Figure 1B, Table supplement 1). Importantly, this increase in the prevalence of infection in 2017 was only observed among the older age groups (i.e., ≥ 6 years) (Figure 1C, Table supplement 1). This difference by age group in 2017 can be attributed to SMC, which only targets children between 3-59 months (i.e., < 5 years). A notable increase in parasite prevalence for adolescents (11-20 years) and adults (> 20 years) was found in 2017 relative to the 2012 (pre-IRS) (Figure 1C, Table supplement 1).

Next, we wanted to explore changes in population size measured by MOIvar. As this metric is based on non-overlap of var repertoire diversity of individual isolates, specifically non-upsA DBLα types, we investigated whether DBLα isolate repertoire similarity (or overlap), as measured by PTS, increased following the sequential interventions (i.e., IRS and SMC). Figure 2 shows that median PTS values for both upsA and non-upsA DBLα types remained low in all surveys, although the PTS distributions for both groups changed significantly at each of the study time points relative to the 2012 baseline survey (pre-IRS) (p-values < 0.001, Kruskal-Wallis test) (Figure 2). Somewhat unexpectedly, the change was in the direction of reduced similarity (i.e., less overlap) with lower median PTS scores and a larger number of isolates sharing no DBLα types (i.e., PTS = 0) in 2014, 2015, and 2017 compared to 2012. Relevant to measurement of MOIvar, the median PTS scores for non-upsA DBLα types were lower following the IRS intervention (PTSnon-upsA: 2014 = 0.013 and 2015 = 0.013 vs. PTSnon-upsA: 2012 = 0.020). In 2017, the non-upsA PTS distributions shifted back towards higher median PTS scores (PTSnon-upsA = 0.016) and fewer isolates shared no DBLα types relative to 2014 and 2015 (Figure 2). To verify this pattern was not influenced by multiclonal infections (MOIvar > 1) we also examined isolates with monoclonal infections (MOIvar = 1) and found that this non-overlapping structure persisted regardless of infection complexity, particularly for the non-upsA DBLα types (Figure supplement 4). These PTS data make clear that we were dealing with a large, highly diverse parasite population where 99.9% of the isolate comparisons in all surveys had PTSnon-upsA scores ≤ 0.1 (i.e., shared ≤ 10% of their non-upsA DBLα types), indicating that DBLα isolate repertoires were highly unrelated (Figure 2). In fact, throughout the IRS, SMC, and subsequent rebound, very few DBLα isolate repertoires were observed to be related, with < 0.003% isolate comparisons in each survey having a PTSnon-upsA ≥ 0.5 (i.e., siblings or recent recombinants) (Figure 2).

Sharing of upsA and non-upsA DBLα types among the DBLα isolate repertoires in 2012 (pre-IRS, red), 2014 (during IRS, orange), 2015 (post-IRS, green), and 2017 (SMC, purple).

Earlier version: The overlapping density and violin plots (upper right-hand corners) show the distribution of PTS scores (i.e., DBLα isolate repertoire similarity) between the (A) upsA and (B) non-upsA DBLα isolate repertoires for those isolates with DBLα sequencing data (Table supplement 2) in each survey. The PTS scales in the density plots have been zoomed-in to provide better visualization of the upsA and non-upsA DBLα types PTS distributions. The colored dashed lines in the density plots indicate the median PTS scores in each survey for the upsA (2012 (red) = 0.078, 2014 (orange) = 0.063, 2015 (green) = 0.054, and 2017 (purple) = 0.064) and non-upsA (2012 (red) = 0.020, 2014 (orange) = 0.013, 2015 (green) = 0.013, and 2017 (purple) = 0.016) DBLα types. Note: The non-upsA median PTS values in 2014 (orange) and 2015 (green) were both 0.013 and overlap in the figure. In the PTS violin plots the central box plots indicate the medians (centre line), interquartile ranges (IQR, upper and lower quartiles), whiskers (1.5x IQR), and outliers (points).

Revised version: The overlapping density and violin plots (upper right-hand corners) show the distribution of PTS scores (i.e., DBLα isolate repertoire similarity) between the (A) upsA and (B) non-upsA DBLα isolate repertoires for those isolates with DBLα sequencing data (Table supplement 2) in each survey. The PTS scales in the density plots have been zoomed-in to provide better visualization of the upsA and non-upsA DBLα types PTS distributions. The coloured dashed lines in the density plots indicate the median PTS scores in each survey for the upsA (2012 (red) = 0.078, 2014 (orange) = 0.063, 2015 (green) = 0.054, and 2017 (purple) = 0.064) and non-upsA (2012 (red) = 0.020, 2014 (orange) = 0.013, 2015 (green) = 0.013, and 2017 (purple) = 0.016) DBLα types. Note: The non-upsA median PTS values in 2014 (orange) and 2015 (green) were both 0.013 and overlap in the figure. In the PTS violin plots the central box plots indicate the medians (centre line), interquartile ranges (IQR, upper and lower quartiles), whiskers (1.5x IQR), and outliers (points).

The raw data of non-upsA DBLα isolate repertoire sizes (Figure supplement 5) were used to estimate MOIvar as adjusted using the Bayesian approach based on pooling the maximum a posteriori MOI estimates (Figure 3) and the mixture distribution (Figure supplement 6). We observed that at baseline in 2012, the majority (89.2%) of the population across all ages carried multiclonal infections (median MOIvar = 4 [IQR: 2 – 6]) (Figure 3A). Following the IRS intervention, the estimated MOIvar distributions became more positively skewed, indicating that a lower proportion of participants harboured multiclonal infections with a lower median MOIvar in 2014 (64.5%; median MOIvar = 2 [IQR: 1 – 3]) and 2015 (71.4%; median MOIvar = 2 [IQR: 1 – 3]) compared to 2012 (Figure 3A). These reductions in median MOIvar and the proportion of multiclonal infections, which were observed across all age groups (Figure 3B), are consistent with the > 90% decrease in transmission intensity following the IRS in turn reducing exposure to new parasite genomes. However, in 2017, both median MOIvar (3 [IQR: 2 – 4]) and the proportion of multiclonal infections (78.9%) rebounded in all age groups, even among the younger children (1-5 years) predominantly targeted by SMC (Figure 3). While the prevalence of infection in 2017 remained low for the younger children (1-5 years), those infected still carried multiclonal infections (84.1% of those infected) (Figure 3B). Although the MOIvar distributions across all age groups started to rebound in 2017 (i.e., less positively skewed compared to 2014 and 2015) they had not fully recovered to the 2012 baseline patterns (Figure 3). This was most apparent among the younger children (1-5 years), as a larger proportion of isolates in 2017, compared to 2012, had MOIvar values equal to one or two, while a smaller proportion had MOIvar values ≥ 5 (Figure 3B).

MOIvar distributions in 2012 (pre-IRS, red), 2014 (during IRS, orange), 2015 (post-IRS, green), and 2017 (SMC, purple) based on pooling the maximum a posteriori MOI estimates.

Earlier version: Estimated MOIvar distributions for the (A) study population and (B) for all age groups (years) in each survey for those isolates with DBLα sequencing data (Table supplements 2 and 7). The median MOIvar values are indicated with the black dashed lines and have been provided in the top right corner (median MOIvar value [interquartile range, upper and lower quartiles]) along with the percentage of P. falciparum infections that were multiclonal (MOIvar > 1) in each survey and age group (years).

Revised version: Estimated MOIvardistributions for the (A) study population and (B) for all age groups (years) in each survey for those isolates with DBLα sequencing data (Table supplements 2 and 7). The median MOIvar values are indicated with the black dashed lines and have been provided in the top right corner (median MOIvar value [interquartile range, upper and lower quartiles]) along with the percentage of P. falciparum infections that were multiclonal (MOIvar > 1) in each survey and age group (years).

Earlier version: Census population size, measured as the number of P. falciparum var repertoires circulating in the population during each survey, was estimated by summation of isolate MOIvar (Figure 4, Table supplement 7). In 2014 during IRS, this number decreased by 72.5% relative to the 2012 baseline survey (pre-IRS) (Figure 4CE), whereas prevalence decreased by 54.5% (Figure 4CE). Although census population size increased slightly in 2015 relative to 2014 (Figure 4A), there were still 64.6% fewer var repertoires in the population compared to 2012 (Figure 4CE) in comparison to a 42.6% decrease in prevalence (Figure 4CE). Importantly this loss of var repertoires in 2014 and 2015 following the IRS intervention was seen for all age groups (Figure 4B), with the greatest overall reductions (> 85%) being observed for the younger children (1-5 years) (Figure 4DF). However, in 2017, the number of diverse var repertoires in the population rebounded, more than doubling between 2015 and 2017 (Figure 4CE). This increase in the number of var repertoires was seen for all age groups in 2017, except for the younger children (1-5 years) where those up to 59 months were targeted by SMC (Figure 4BD). In fact, the greatest overall increase was observed for the adolescents (11-20 years), where the number of var repertoires in 2017 was ∼1.4 times higher compared to 2012 (Figure 4F). Similar trends in the number var repertoires were also observed for the older children (6-10 years) and adults (>20 years) in 2017, although the rebound was not as striking as that detected for the adolescents.

Revised version: Census population size, measured as the number of P. falciparum var repertoires circulating in the population during each survey, was estimated by summation of isolate MOIvar (see Materials and Methods; Figure 4, Table supplement 7). In 2014 during IRS, this number decreased by 71.4% relative to the 2012 baseline survey (pre-IRS) (Figure 4CE), whereas prevalence decreased by 54.5% (Figure 4CE). Although census population size increased slightly in 2015 relative to 2014 (Figure 4A), there were still 64.4% fewer var repertoires in the population compared to 2012 (Figure 4CE) in comparison to a 42.6% decrease in prevalence (Figure 4CE). Importantly this loss of var repertoires in 2014 and 2015 following the IRS intervention was seen for all age groups (Figure 4B), with the greatest overall reductions (≥ 83.8%) being observed for the younger children (1-5 years) (Figure 4DF). However, in 2017, the number of diverse var repertoires in the population rebounded, more than doubling between 2015 and 2017 (Figure 4CE). This increase in the number of var repertoires was seen for all age groups in 2017, except for the younger children (1-5 years) where those up to 59 months were targeted by SMC (Figure 4BD). In fact, the greatest overall increase was observed for the adolescents (11-20 years) and adults (> 20 years), where the number of var repertoires in 2017 was ∼1.2 times higher compared to 2012 (Figure 4F). Similar trends in the number var repertoires were also observed for the older children (6-10 years) in 2017, although the rebound was not as striking as that detected for the adolescents and adults.

Estimated number and relative change in the number of P. falciparum var repertoires in 2012 (pre-IRS, red), 2014 (during IRS, orange), 2015 (post-IRS, green), and 2017 (SMC, purple).

Earlier version: The estimated number of var repertoires (i.e., census population size) for those isolates with DBLα sequencing data (Table supplements 2 and 7) in the (A) study population and (B) for all age groups (years). The estimated number of var repertoires vs. P. falciparum prevalence for (C) study population and (D) for all age groups (years) (Table supplement 7). The percentage change in P. falciparum prevalence (black dotted line) and the estimated number of var repertoires (black solid line) in 2014, 2015, and 2017 compared to the 2012 baseline survey (red dashed horizontal line at 0% change) for the (E) study population and (F) for all age groups (years). Error bars in (A-D) represent the upper and lower limits of the 95% confidence intervals (95% CIs). The 95% CIs for the number of var repertoires (i.e., census population size) were calculated based on a bootstrap approach. We resampled 10,000 replicates from the original population-level distribution with replacement. Each resampled replicate has the same size as the original sample. We then derive the 95% CI based on the distribution of the resampled replicates. The 95% CIs for P. falciparum prevalence (%) were calculated using the Wald Interval.

Revised version: The estimated number of var repertoires (i.e., census population size) for those isolates with DBLα sequencing data (Table supplements 2 and 7) in the (A) study population and (B) for all age groups (years). The estimated number of var repertoires vs. P. falciparum prevalence for (C) study population and (D) for all age groups (years) (Table supplement 7). The percentage change in P. falciparum prevalence (black dotted line) and the estimated number of var repertoires (black solid line) in 2014, 2015, and 2017 compared to the 2012 baseline survey (red dashed horizontal line at 0% change) for the (E) study population and (F) for all age groups (years). Error bars in (A-D) represent the upper and lower limits of the 95% confidence intervals (95% CIs). To account for differences in sampling depth across age groups and surveys, we performed subsampling with replacement by selecting the minimum number of individuals in each age group across all surveys. We then calculated the total number of var repertoires from these subsampled individuals within each age group in each survey. This approach ensures consistent sample sizes within each age group across all surveys. Finally, we summed the var repertoires across age groups to obtain the total var repertoire count for each survey. The mean (coloured solid points) and 95% CIs for the number of var repertoires were estimated by repeating the subsampling procedure 10,000 times. The CIs were then derived from the distribution of these repeated subsampling replicates. The 95% CIs for P. falciparum prevalence (%) were calculated using the Wald Interval.

As census population size changed considerably during the sequential IRS and SMC interventions, we investigated how the removal or loss of P. falciparum var repertoires and subsequent rebound in 2017 altered DBLα type richness, measured as the number of unique upsA and non-upsA DBLα types in the parasite population in each survey. Richness at baseline in 2012 (pre-IRS) was high with a large number of unique DBLα types (upsA = 2,218; non-upsA = 33,159) (Figure 5, Table supplement 2) and limited overlap of var repertoires (i.e., median PTSnon-upsA ≤ 0.020) seen in a relatively small study population of 685 microscopically positive individuals (Figure 2). In 2014, as P. falciparum prevalence and census population size declined (Figure 4) so too did the number of DBLα types, resulting in a 32.2% and 55.3% reduction in richness for the upsA and non-upsA DBLα types, respectively, compared to 2012 (Figure 5, Table supplement 2). Again in 2015, as P. falciparum prevalence and population size remained low (Figure 4) DBLα type richness was still less than that observed in 2012 (24.6% and 46.0% reduction for upsA and non-upsA DBLα types, respectively) (Figure 5, Table supplement 2). Finally, in 2017, we found that upsA and non-upsA DBLα type richness rebounded relative to 2014 and 2015, coincident with the increase in P. falciparum prevalence and census population size (Figure 4, Figure 5).

UpsA and non-upsA DBLα type richness in 2012 (pre-IRS, red), 2014 (during IRS, orange), 2015 (post-IRS, green), and 2017 (SMC, purple).

Number of unique (A) upsA and (B) non-upsA DBLα types (i.e., richness) observed in each survey vs. P. falciparum prevalence based on those isolates with DBLα sequencing data (Table supplement 2). Error bars represent the upper and lower limits of the 95% confidence intervals (95% CIs) for the P. falciparum prevalence (%; x-axis) and ± 2 standard deviations (± 2SD) for the number of unique upsA and non-upsA DBLα types (y-axis). The 95% CIs for P. falciparum prevalence (%) were calculated using the Wald Interval. The ± 2SD for the number of unique upsA and non-upsA DBLα types were calculated based on a bootstrap approach. We resampled 10,000 replicates from the original population-level distribution with replacement. Each resampled replicate has the same size as the original sample. We then derive the standard deviation (SD) based on the distribution of the resampled replicates.

Earlier version: Given this reduction in DBLα type richness following the IRS intervention and subsequent rebound in 2017, we wanted to explore whether the loss of richness was influenced by the frequency of individual DBLα types in the parasite population within and among surveys. To answer this, we defined the relative frequency of individual DBLα types in all isolates in each survey (Figure 6, Figure supplement 7). We discovered that individual upsA and non-upsA DBLα types were not all at equal frequencies within a survey and among surveys. They could be classified as frequent (i.e., observed in 11-20 or > 20 isolates), less frequent (i.e., observed in 2-10 isolates) or only seen once, at baseline in 2012 (Figure 6AB). In 2014 and 2015, following IRS, there was a significant increase in the proportion of upsA and non-upsA DBLα types in the lower frequency categories (p-value < 0.001, Mann-Whitney U test), with all DBLα types becoming rarer in the population (Figure 6CD). This change can be attributed to the removal of P. falciparum var repertoires (Figure 4) with associated loss of upsA and non-upsA DBLα type richness (Figure 5), which disproportionally affected those DBLα types seen once. This shift to all DBLα types becoming rarer following IRS changed in 2017, where the proportion of DBLα types in the more frequent categories (i.e., 2-10, 11-20, or > 20 isolates) significantly increased while the proportion seen once decreased (p-values < 0.001, Mann-Whitney U tests) (Figure 6CD). Data in Figure 6 (A-D) pointed to a differential effect of the IRS intervention and subsequent rebound on the less frequent upsA and non-upsA DBLα types vs. those that were classified as frequent, where those DBLα types that were most frequent persisted longitudinally.

Revised version: Given this reduction in DBLα type richness following the IRS intervention and subsequent rebound in 2017, we wanted to explore whether the loss of richness was influenced by the frequency of individual DBLα types in the parasite population within and among surveys. To answer this, we defined the relative frequency of individual DBLα types in all isolates in each survey (Figure 6 and Figure supplement 7). We discovered that individual upsA and non-upsA DBLα types were not all at equal frequencies within a survey and among surveys. They could be classified as frequent (i.e., observed in 11-20 or > 20 isolates), less frequent (i.e., observed in 2-10 isolates) or only seen once, at baseline in 2012 (Figure 6AB). In 2014 and 2015, following IRS, there was a significant increase in the proportion of upsA and non-upsA DBLα types in the lower frequency categories (p-value < 0.001, Mann-Whitney U test), with all DBLα types becoming rarer in the population (Figure 6CD). This change can be attributed to the removal of P. falciparum var repertoires (Figure 4) with associated loss of upsA and non-upsA DBLα type richness (Figure 5), which disproportionally affected those DBLα types seen once. This shift to all DBLα types becoming rarer following IRS changed in 2017, where the proportion of DBLα types in the more frequent categories (i.e., 2-10, 11-20, or > 20 isolates) significantly increased while the proportion seen once decreased (p-values < 0.001, Mann-Whitney U tests) (Figure 6CD). Data in Figure 6 (A-D) pointed to a differential effect of the IRS intervention and subsequent rebound on the less frequent upsA and non-upsA DBLα types vs. those that were classified as frequent, where those DBLα types that were most frequent persisted longitudinally.

UpsA and non-upsA DBLα type frequencies and survival in 2012 (pre-IRS, red), 2014 (during IRS, orange), 2015 (post-IRS, green), and 2017 (SMC, purple).

Earlier version: Heatmaps showing the patterns of diversity for the (A) upsA and (B) non-upsA DBLα types. The columns represent all the upsA and non-upsA DBLα types observed in the four surveys, and the rows represent each of the 2,802 upsA DBLα types and the 50,436 non-upsA DBLα types (Table supplement 2). White rows are used to denote the absence of a specific DBLα type, while the presence of a DBLα type is indicated in colour and further categorised (colour gradations) based on the frequency or the number of times (i.e., number of isolates) a DBLα type was observed in each survey (Frequency categories: 1, 2-10, 11-20, > 20 isolates; Note the frequency category cut-offs were chosen based on the frequency distributions in Figure supplement 7). The proportions of (C) upsA and (D) non-upsA DBLα types in each survey based on the number of times (i.e., number of isolates) they were observed in each survey. Kaplan-Meier survival curves for the (E) upsA and (F) non-upsA DBLα types across time (2012 to 2017) categorised based on their frequency at baseline in 2012 (pre-IRS, red). The colored shaded areas represent the upper and lower limits of the 95% confidence intervals (95% CIs) with the number (N) of upsA and non-upsA DBLα types in each frequency category are provided in parenthesis. These survival curves include only those upsA (N = 2,218) and non-upsA (N = 33,159) DBLα types that were seen at baseline in 2012 (pre-IRS) as indicated in red (Table supplement 2). The x-axis indicates time where time “0” denotes 2012 (pre-IRS), “1” denotes 2014 (during IRS), “2” denotes 2015 (post-IRS), and finally “3” denotes 2017 (SMC). Note: In the survival curves the 11-20 and > 20 frequency categories for both the (E) upsA and (F) non-upsA DBLα types overlap in the figure.

Revised version: Heatmaps showing the patterns of diversity for the (A) upsA and (B) non-upsA DBLα types. The columns represent all the upsA and non-upsA DBLα types observed in the four surveys, and the rows represent each of the 2,802 upsA DBLα types and the 50,436 non-upsA DBLα types (Table supplement 2). White rows are used to denote the absence of a specific DBLα type, while the presence of a DBLα type is indicated in colour and further categorised (colour gradations) based on the frequency or the number of times (i.e., number of isolates) a DBLα type was observed in each survey (Frequency categories: 1, 2-10, 11-20, > 20 isolates; Note the frequency category cut-offs were chosen based on the frequency distributions in Figure supplement 7). The proportions of (C) upsA and (D) non-upsA DBLα types in each survey based on the number of times (i.e., number of isolates) they were observed in each survey. Kaplan-Meier survival curves for the (E) upsA and (F) non-upsA DBLα types across time (2012 to 2017) categorised based on their frequency at baseline in 2012 (pre-IRS, red). The coloured shaded areas represent the upper and lower limits of the 95% confidence intervals (95% CIs) with the number (N) of upsA and non-upsA DBLα types in each frequency category are provided in parenthesis. These survival curves include only those upsA (N = 2,218) and non-upsA (N = 33,159) DBLα types that were seen at baseline in 2012 (pre-IRS) as indicated in red (Table supplement 2). The x-axis indicates time where time “0” denotes 2012 (pre-IRS), “1” denotes 2014 (during IRS), “2” denotes 2015 (post-IRS), and finally “3” denotes 2017 (SMC). Note: In the survival curves the 11-20 and > 20 frequency categories for both the (E) upsA and (F) non-upsA DBLα types overlap in the figure.

To explore this observation further we restricted the longitudinal analysis to those DBLα types from the baseline survey in 2012 (pre-IRS). We compared the probability of survival for the DBLα types identified at baseline in 2012 and found that the upsA DBLα types persisted significantly longer in the population relative to the non-upsA DBLα types (p < 0.001, log-rank test), despite the IRS intervention. The simple explanation being that although the upsA DBLα types had lower richness (Figure 5), a larger proportion was classified as frequent indicating that multiple copies existed in the population compared to the non-upsA DBLα types (Figure 6CD). Furthermore, when we examined survival using the frequency categories, the upsA and non-upsA DBLα types that were observed at multiple study time points (i.e., 2012, 2014, 2015, and 2017), albeit in different isolate repertoires, were those that were most frequent (i.e., observed in 11-20 and >20 isolates) in the population at baseline in 2012 (Figure 6EF). As expected, the DBLα types that were only observed once in 2012, were significantly less likely to be seen longitudinally (p-value < 0.001, log-rank test) (Figure 6EF). These differential changes in DBLα type richness with respect to rare vs. frequent DBLα types are a consequence of changes in census population size with interventions (Figure 4) where each isolate repertoire is composed of many rare DBLα types as defined by PTS (Figure 2).

Discussion

P. falciparum populations in high-transmission endemic areas in SSA, are characterised by extensive diversity, high rates of recombination, as well as frequent multiclonal infections. Here we defined census population size of P. falciparum to understand total parasite diversity in a human population and explore the age-specific efficacy of malaria interventions to reduce this metric in such areas as typified by Bongo, Ghana. Census population size proved more informative than parasite prevalence alone because it captures “within” host parasite population size as MOI, rather than using the infected host per se as a unit of population size. Whilst the concept of census population size is agnostic of how you measure MOI, the extensive DBLα isolate repertoire diversity presented makes a strong case for fingerprinting parasite isolates by varcoding in high transmission. This is opposed to looking at allelic diversity of a single copy antigen gene, such as csp.

Earlier version: Census population size is a total enumeration or count of infections in a given population sample and over a given time period in an ecological sense, distinct from the formal effective population size (Ne) used in population genetics (Charlesworth, 2009). Given the low overlap between var repertoires of parasites observed in monoclonal infections (MOI = 1), the census population size calculated in Bongo, Ghana translates to a diversity of strains or repertoires. The distinction of census population size in terms of infection counts and effective population size from population genetics has been made before for pathogens, including the seasonal influenza virus and the measles virus (Bedford et al., 2011); it is also a distinction made in the ecological literature for non-pathogen populations (Palstra and Fraser, 2012). The census population size of a given population sample depends of course on sample size and was used here for comparisons across time of samples of the same depth (i.e., ∼2,000 individuals). There is, however, a simple map between census population size and mean MOI, as one can simply divide or multiply by the sample size, respectively, to convert between the two quantities (Table supplement 7). Therefore, one can extrapolate from the census population size of a given population sample to that of the whole population of local hosts to compare across studies that differ in sampling depth. What is needed for this extrapolation is a stable mean MOI relative to the sample size or sampling depth, which is indeed the case in this study (Figure supplement 8) and can be easily checked in other studies. Given the typical duration of infection, we expect our population size to be representative of a per-generation measure.

Revised version: Census population size is a total enumeration or count of infections in a given population sample and over a given time period in an ecological sense, distinct from the formal effective population size (Ne) used in population genetics (Charlesworth, 2009). Given the low overlap between var repertoires of parasites observed in monoclonal infections (MOI = 1), the census population size calculated in Bongo, Ghana translates to a diversity of strains or repertoires. The distinction of census population size in terms of infection counts and effective population size from population genetics has been made before for pathogens, including the seasonal influenza virus and the measles virus (Bedford et al., 2011); it is also a distinction made in the ecological literature for non-pathogen populations (Palstra and Fraser, 2012). The census population size of a given population sample depends of course on sample size and was used here for comparisons across time of samples of the same depth (i.e., ∼2,000 individuals). There is, however, a simple map between census population size and mean MOI, as one can simply divide or multiply by the sample size, respectively, to convert between the two quantities (Table supplement 7). Therefore, one can extrapolate from the census population size of a given population sample to that of the whole population of local hosts in a given area to compare across studies that differ in sampling depth and/or spatial extent. What is needed for this extrapolation is a stable mean MOI relative to the sample size or sampling depth, which is indeed the case in this study (Figure supplement 8) and can be easily checked in other studies. Given the typical duration of infection, we expect our population size to be representative of a per-generation measure.

By varcoding we identified a very large census parasite population size in a relatively small human population of ∼2,000 individuals at baseline and captured age-specific changes in this metric in response to sequential malaria control interventions. IRS reduced the MOIvar parasite population size substantially with the greatest reductions (85%) seen in the younger children (1-5 years). More than two years after cessation of IRS, rebound in 2017 was rapid in all age groups, except for the younger children (1-5 years) where those up to 59 months were targeted by SMC. Population sizes in adolescents (11-20 years) and adults (> 20 years) showed they carried more infections in 2017 than at baseline in 2012. This is indicative of a loss of immunity during IRS which may relate to the observed loss of var richness, especially the many rare types. This warrants further investigation of changes in variant-specific immunity. During and following the IRS and SMC interventions, var diversity remained high and var repertoire overlap remained low, reflecting characteristic properties of high transmission and demonstrating the overall resilience of the species to survive significant short-term perturbations. Combining interventions and targeting older age groups or the whole community with chemoprevention would no doubt have a much greater impact on reducing the diversity of the reservoir of infection.

What was striking about the Bongo study was the speed with which rebound in MOIvar per person and census population size occurred, once the short-term IRS was discontinued. We looked for a potential explanation in our genetic data. PTS and population frequency data showed that many of the DBLα types occurred in multiple repertoires or genomes. This enabled the survival of these more frequent DBLα types through the interventions, facilitating rebound by maintenance of this diversity. The other notable population genetic result of our study was the failure to increase similarity (or relatedness by state) of var repertoires by reducing transmission by > 90% via IRS. From a baseline of a very large population size with very low overlap in repertoires you need outcrossing to create relatedness. However, this was less likely to happen due to reduced transmission as a result of IRS. Rebound, with associated increases in transmission, led to a small increase in var repertoire similarity. This is the opposite to what has been observed in areas of lower transmission under intense malaria control where the intensity of interventions led to increased genome similarity, as assessed by IBD, from a starting point of much lower genome diversity and greater relatedness (Daniels et al., 2015).

Our molecular approach to measure population size has been to sum MOIvar in individual hosts with microscopically-detectable infections. Like any diagnostic method there are limits to sensitivity and specificity, which can be more or less tolerated dependent upon the purpose of the study. Here we have looked at relative changes in population size with sequential interventions using an interrupted time-series study design and observed changes by measuring MOIvar. We have accounted for missing DBLα type data where complete var repertoires may not have been sequenced using a Bayesian method based on empirical knowledge of the measurement error. This approach has conceptual relation to the Bayesian approach by Johnson and Larremore (Johnson and Larremore, 2022) to estimate complete repertoire size of, and overlap between, monoclonal infections from incomplete sampling of DBLα types. Our measurement of population size based on MOIvar will be subject to other sampling errors which may in the end be more significant (discussed in detail in Labbé et al. (Labbé et al., 2023)). For example, low parasitemia typical of asymptomatic infections, small blood volumes, clinical status, and/or within host dynamics, including synchronicity, will all create sampling problems, but these are common to all measures of MOI.

Previously, we have drawn attention to the potential underestimation of the number of DBLα types of related parasites generated by a cross, when using var genotyping (Labbé et al., 2023). Such related parasites must be created frequently in high transmission due to extensive outcrossing (Babiker et al., 1994; Paul et al., 1995). Single clone genomics experiments using biallelic SNPs from whole genome sequencing data have also detected related parasites using IBD in clinical infections from humans in a high-transmission area of Malawi (Nkhoma et al., 2020). Here we have analysed low-to moderate-density, chronic, asymptomatic infections (see Table supplement 1) under strong immune selection in semi-immune hosts whom we have shown select against parasites with high PTS scores consistent with relatedness by decent (Day et al., 2017; He et al., 2018; Ruybal-Pesántez et al., 2022). When considering the importance of possible exclusion of parasites related by descent, the only sure way to detect such parasites in high transmission is by single cell genomics, a methodology of limited application to malaria surveillance due to practicality and cost of scale up. Again, the error from failure to sample infections related by descent must be weighed up against the issues of under-sampling as described above.

The Bayesian approach of the varcoding method relies on the low or limiting similarity of var repertoires infecting individual human hosts. As such, it would appear to break down as the var repertoire overlap moves away from extremely low, and therefore, for locations with lower transmission intensity. Interestingly, this is not the case in the numerical simulations of Labbé et al. (Labbé et al., 2023), for a gradient of three transmission intensities, from high to low, with the original varcoding method performing well across the gradient. This robustness of the method may arise from a nonlinear and fast transition from low to high overlap that is accompanied by MOI changing rapidly from primarily multiclonal (MOI > 1) to monoclonal (MOI = 1) infections. This matter needs to be investigated further in the future, including ways to extend the Bayesian approach to explicitly include the distribution of var repertoire overlap.

In summary, our findings provide parasite population insights into why rebound is the inevitable consequence of such short-term IRS interventions unless you simultaneously target the highly diverse, long-lived parasite population in humans, not just children < 5 years by SMC. Of potential translational significance for malaria molecular surveillance, we identify new metrics, especially MOIvar and census population size as well as var frequency category, informative to monitor and evaluate interventions in high-transmission areas with multiclonal infections and high rates of outcrossing. Such metrics could be used longitudinally to detect incremental gains of transmission-reducing interventions, including IRS, LLINs, and vaccines to perturb the high-transmission characteristics of the parasite population in humans in high-burden countries in SSA.